```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment03_COMM187W26.ipynb")
```

# COMM 187 W26 Coding Assignment #3

<span style="color:red">**!!! PLEASE MAKE SURE TO RUN THE FIRST CELL ABOVE WHICH SAYS "# Initialize Otter" BEFORE RUNNING AUTO-GRADER !!!**</span>

## <center>TWEETS BY US SENATORS (2008-17)</center>

![](./imgs/A3_introimg.png)

The dataset for this assignment is about the favorites (likes), replies, and retweets on tweets made by US Senators between 2008 and 2017, as recorded on Oct. 19 and 20, 2017. This dataset was collected and analyzed for the following [blog post by FiveThirtyEight](https://fivethirtyeight.com/features/the-worst-tweeter-in-politics-isnt-trump/), and you can read more information about dataset [here](https://github.com/fivethirtyeight/data/tree/master/twitter-ratio/).

The original dataset for this article contains over 288 thousand rows. For the purpose of this assignment, you have beeen provided with a reduced dataset of randomly selected five thousand rows, which is stored in the "data" subfolder and is named `senators_tweets.csv`. It contains the following columns:


Header | Definition
--- | ---
`created_at` | Date and time of the creation of the tweet in the format "mm/dd/yy hh:mm"
`text` | The text content of the tweet
`url` | URL of the tweet
`replies` | Number of replies for the tweet when data was collected
`retweets` | Number of retweets for the tweet when data was collected
`favorites` | Number of favorites (or likes) for the tweet when data was collected
`user` | Twitter username of the US Senator
`bioguide_id` | Unique ID for each US Senator
`party` | Party identification for each US Senator: 'D' for Democrat, 'R' for Republican, and 'I' for Independent
`state` | The state being represented by the US Senator

---

```python
# RUN THIS CELL BEFORE ATTEMPTING THE ASSIGNMENT QUESTIONS
import numpy as np
import pandas as pd
from scipy import stats
```

---

## Question 1: Loading Data (0.25 point)

Use the pandas `read_csv` function to load the data `senator_tweets.csv` file into a pandas DataFrame named `df`.

```python
### Write your code below (in place of ...)
df = ...
```

```python
grader.check("q1")
```

---

## Question 2: Senator Jack Reed's Tweets (1.25 points)

Calculate the following for the tweets made by Senator Jack Reed (user ID = "SenJackReed"):

 - Mean number of retweets, store it in `mean_rt_JackReed` **(0.25 point)**
 - Standard deviation of retweets, using pandas `.std()` method, store it in `std_rt_JackReed` **(0.25 point)**
 - Median number of replies, store it in `median_replies_JackReed` **(0.25 point)**
 - Standard deviation of replies, using numpy `np.std()` function, store it in `std_replies_JackReed` **(0.25 point)**
 - Mean number of favorites, store it in `mean_fav_JackReed` **(0.25 point)**

```python
### Write your code below (in place of ...)
mean_rt_JackReed = ...
std_rt_JackReed = ...
median_replies_JackReed = ...
std_replies_JackReed = ...
mean_fav_JackReed = df[df['user'] == "SenJackReed"]['favorites'].mean()
```

```python
grader.check("q2")
```

---

## Question 3: Comparing Favorites on Tweets Across Parties

### A. (1 point)

Compare the **favorites** (or *likes*) across the three parties (stored in the column `party` with values **D**, **R**, and **I**) using ANOVA. \
Store the F statistic in `q3a_f` and the p-value in `q3a_p`.

**Steps:**
 - Extract `favorites` column for each party (D, R, and I) and store them in respective variables `D`, `R`, and `I`
 - Calculate the ANOVA using `stats.f_oneway()` function
 - **Tip:** by writing `var1, var2 = stats.f_oneway()`, you can store the F statistic in `var1` and p value in `var2`

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q3a")
```

### B. (0.5 point)

Compare the number of **favorites** (or *likes*) between the two major parties (D and R) using t-test. Store the t-statistic in `q3b_t` and the p-value in `q3b_p`.

Steps:
 - Use the variables `D` and `R` defined in the previous question
 - Calculate t-test using `stats.ttest_ind()` function
 - **Tip:** by writing `var1, var2 = stats.ttest_ind()`, you can store the T statistic in `var1` and p value in `var2`

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q3b")
```

---

## Question 4: Comparing retweets across states (1 point)

Compare the **retweets** across the senators from the following six states using ANOVA:
 - South Carolina (SC)
 - Arizona (AR)
 - Oregon (OR)
 - Kansas (KS)
 - New Jersey (NJ)
 - Delaware (DE)

Store the F statistic in `q4_f` and the p-value in `q4_p`. 

Follow similar steps as in Question 3A. You are not being provided with a list of explicit steps for this question, but you have had enough practice to make this work. You got this :)


```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q4")
```

---

Tests for q1 (points: 0.25)
```
>>> all(df == pd.read_csv('./data/senator_tweets.csv'))
True
```
```
>>> type(df) == pd.core.frame.DataFrame
True
```


Tests for q2 (points: 1.25)
```
>>> bool(mean_rt_JackReed == df[df['user'] == 'SenJackReed']['retweets'].mean())
True
```
```
>>> bool(std_rt_JackReed == df[df['user'] == 'SenJackReed']['retweets'].std())
True
```
```
>>> bool(median_replies_JackReed == df[df['user'] == 'SenJackReed']['replies'].median())
True
```
```
>>> bool(std_replies_JackReed == np.std(df[df['user'] == 'SenJackReed']['replies']))
True
```
```
>>> bool(mean_fav_JackReed == df[df['user'] == 'SenJackReed']['favorites'].mean())
True
```


Tests for q3a (points: 1)
```
>>> bool(abs(q3a_f - 12.434260538101874) < 0.0001)
True
```
```
>>> bool(abs(q3a_p - 0) < 1e-05)
True
```


Tests for q3b (points: 0.5)
```
>>> bool(abs(q3b_t - 2.389643559096571) < 0.0001)
True
```
```
>>> bool(abs(q3b_p - 0.016902486599381032) < 0.0001)
True
```


Tests for q4 (points: 1)
```
>>> bool(abs(q4_f - 1.2590702223944095) < 0.0001)
True
```
```
>>> bool(abs(q4_p - 0.27975665692284135) < 0.0001)
True
```
