```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment01_COMM187W26.ipynb")
```

# COMM 187 W26 Coding Assignment #1

<span style="color:red">**!!! PLEASE MAKE SURE TO RUN THE FIRST CELL ABOVE WHICH SAYS "# Initialize Otter" BEFORE RUNNING AUTO-GRADER !!!**</span>

### Question 1

**A. (0.25 point)** 

A communication researcher is analyzing the time spent by one individual on *Instagram*. They collected data on the time spent (in minutes) on Instagram for 7 consecutive days. The times spent on each day are stored in variables `ig_day_01` through `ig_day_07`, as shown below.

Create a Python **list** that contains all seven daily Instagram usage times in sequence from day 1 to day 7. Store this list in a new variable called `ig_list`. Remember that in Python, lists are created using square brackets `[ ]` with items separated by commas.

```python
ig_day_01 = 120
ig_day_02 = 110
ig_day_03 = 15
ig_day_04 = 130
ig_day_05 = 135
ig_day_06 = 100
ig_day_07 = 120

### Write your code below (in place of ...)
...
```

```python
grader.check("q1a")
```

**B. (0.25 point)** 

Immediately following the 7 consecutive days, the communication researcher continues to observe the same individual's Instagram usage for four more days (days 8 through day 11).

The data collected shows:
- Day 8: 90 minutes
- Day 9: 190 minutes
- Day 10: 145 minutes
- Day 11: 30 minutes

Create a new list containing just these four new daily values in sequence (day 8 through day 11). Store this list in a variable called `ig_list_later`.

```python
### Write your code below (in place of ...)
ig_list_later = ...
```

```python
grader.check("q1b")
```

**C. (0.5 point)** 

Now combine the two lists you created (`ig_list` and `ig_list_later`) to make a single list containing all Instagram usage data from day 1 through day 11 in sequence.

Store your result in a new variable called `ig_list_updated`.

**Important:** Use Python's list operations to combine `ig_list` and `ig_list_later` rather than manually typing out all the numbers again. This programming approach allows you to automatically update the combined list whenever the original lists change.

```python
### Write your code below (in place of ...)
ig_list_updated = ...
```

```python
grader.check("q1c")
```

**D. (0.5 point)** 

A new researcher needs only the data from days 2 through 9 (inclusive of day 9).

Create a new variable called `ig_list_subset` containing just this portion of the data from your `ig_list_updated` variable.

**Important:** 
- Use Python's indexing on `ig_list_updated` rather than typing the numbers manually
- Remember that in Python, list indexing starts at 0 (so day 1 is at index 0, day 2 is at index 1, etc.)
- In list slicing in Python, the notation is `list[start:stop]` where the start index is included but the stop index is excluded

```python
### Write your code below (in place of ...)
ig_list_subset = ...
```

```python
grader.check("q1d")
```

### Question 2

**A. (0.5 point)**

You have data on the number of texts sent by **Ace** to his friends **Buck** and **Chip** in a given week.

Ace's texts to Buck:
 - Sunday: 10 texts
 - Monday: 3 texts
 - Tuesday: 3 texts
 - Wednesday: 47 texts
 - Thursday: 5 texts
 - Friday: 0 texts
 - Saturday: 1 text

Ace's texts to Chip:
 - Sunday: 1 text
 - Monday: 5 texts
 - Tuesday: 0 texts
 - Wednesday: 4 texts
 - Thursday: 52 texts
 - Friday: 21 texts
 - Saturday: 10 texts

Create the following:
1. A numpy array called `buck` with the number of texts from Ace to Buck in sequence from Sunday to Saturday
2. A numpy array called `chip` with the number of texts from Ace to Chip in sequence from Sunday to Saturday

```python
import numpy as np

### Write your code below (in place of ...)
buck = ...
chip = ...
```

```python
grader.check("q2a")
```

**B. (0.5 point)** 

Turns out that Ace texted a third friend, **Domingo**, this week:

Ace's texts to Domingo:
- Sunday: 0 texts
- Monday: 0 texts
- Tuesday: 92 texts
- Wednesday: 0 texts
- Thursday: 1 text
- Friday: 0 texts
- Saturday: 2 texts

Complete these tasks:
1. Create a numpy array called `domingo` with the number of texts from Ace to Domingo in sequence from Sunday to Saturday
2. Calculate the total texts sent by Ace each day by adding the arrays for all three friends (Buck, Chip, and Domingo)
3. Store this total in a numpy array called `ace_total`

Note: You can use the `buck` and `chip` arrays you created in part A.

```python
### Write your code below (in place of ...)
domingo = ...
ace_total = ...
```

```python
grader.check("q2b")
```

### Question 3

You are conducting digital marketing research for a beauty brand that is advertising on Instagram. 

**Consumer Query Response Time (CQRT)** is the average number of minutes that this brand takes to respond to consumer queries in Instagram DMs. 

You have data on their daily CQRT for the past 120 days, stored in the variable `brand_cqrt`.

**A. (0.5 point)**

Using the NumPy function `np.mean()`, calculate the **mean (average) CQRT** for this brand across all 120 days. \
Store this result in a new variable called `brand_cqrt_mean`.

```python
brand_cqrt = np.array([ 74,  67,  16,  56,  40,  49,  12, 113,  57, 104, 110, 102,  54,
                       70,  53, 101, 109,  22,  46,  21,  37,  12,  67,  15, 109,  25,
                       72, 115,  80,   6, 100,  19, 115, 100,  71,  17, 110,  56,  93,
                       46,  21,  94,  98,  39,  12,   5, 101, 104,  18,  77,  42,  91, 
                       101, 110,  45,  85,  61,  88,  31,  58,  71,  37,   9,  46,  24,
                       1, 100,  75,  73,  13,   1,  13,  92,  71,  66,  82,  17,  21, 
                       92,  80,   6,  92,  94,  81, 105, 103, 105,  56,  44,  28,  10, 
                       80, 111,  85,   4, 103, 107, 105,  74,  78,  86,  56,  49,  28, 
                       105,  37,  99,  30, 100,  34,  61,   2,  21,  60,  35,  72,  60, 38,  31,  18])

### Write your code below (in place of ...)
brand_cqrt_mean = ...
```

```python
grader.check("q3a")
```

**B. (0.5 point)**

Using the NumPy function `np.median()`, calculate the **median CQRT** for this brand in the past 120 days.

Store this result in a new variable called `brand_cqrt_median`.

Note: While the mean gives you the average value, the median gives you the middle value when all response times are arranged in order (and is less affected by extreme values).

```python
### Write your code below (in place of ...)
brand_cqrt_median = ...
```

```python
grader.check("q3b")
```

**C. (0.5 point)**

You also have data on the number of consumers the brand interacted with on Instagram for each of the past 120 days, stored in the variable `brand_consumers`.

To find the total time spent on customer interactions each day:

 - Calculate the total minutes spent by multiplying the number of consumers each day (`brand_consumers`) by the average response time for that day (`brand_cqrt`)
 - Store this result in a new variable called `brand_total_minutes`

Note: This multiplication works because if the brand spent an average of X minutes per consumer on a given day and interacted with Y consumers that day, then the total time spent was X times Y minutes.

```python
brand_consumers = np.array([45, 20, 16, 42, 33, 13, 29, 33, 15, 22, 14, 46, 24, 32, 15, 12, 48, 
                            35, 12, 42, 24, 34, 45, 36, 23, 41, 25, 45, 36, 14, 20, 19, 39, 33,
                            53, 42, 22, 45, 24, 52, 12, 30, 10, 34, 10, 45, 43, 34, 29, 12, 42,
                            30, 21, 33, 38, 10, 42, 25, 50, 32, 17, 35, 35, 36, 27, 24, 15, 18,
                            49, 45, 32, 43, 53, 32, 46, 52, 22, 12, 32, 43, 16, 39, 26, 33, 11,
                            22, 42, 53, 38, 35, 19, 44, 41, 34, 36, 46, 29, 40, 13, 23, 51, 28,
                            31, 18, 28, 10, 12, 11, 48, 43, 47, 43, 53, 41, 15, 33, 37, 36, 42, 18])

### Write your code below (in place of ...)
brand_total_minutes = ...
```

```python
grader.check("q3c")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Coding Assignment #1 for COMM187: Data Science in Communication Research. Once finished answering questions, first **SAVE** then download this .ipynb file. Submit the file as instructed on Canvas and Gradescope. **ONLY** submit the .ipynb file, NOT the pdf.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1a (points: 0.25)
```
>>> len(ig_list) == 7
True
```
```
>>> ig_list == [120, 110, 15, 130, 135, 100, 120]
True
```


Tests for q1b (points: 0.25)
```
>>> len(ig_list_later) == 4
True
```
```
>>> ig_list_later == [90, 190, 145, 30]
True
```


Tests for q1c (points: 0.5)
```
>>> len(ig_list_updated) == 11
True
```
```
>>> ig_list_updated == [120, 110, 15, 130, 135, 100, 120, 90, 190, 145, 30]
True
```


Tests for q1d (points: 0.5)
```
>>> len(ig_list_subset) == 8
True
```
```
>>> ig_list_subset == [110, 15, 130, 135, 100, 120, 90, 190]
True
```


Tests for q2a (points: 0.5)
```
>>> np.array_equal(buck, np.array([10, 3, 3, 47, 5, 0, 1]))
True
```
```
>>> np.array_equal(chip, np.array([1, 5, 0, 4, 52, 21, 10]))
True
```


Tests for q2b (points: 0.5)
```
>>> np.array_equal(domingo, np.array([0, 0, 92, 0, 1, 0, 2]))
True
```
```
>>> np.array_equal(ace_total, np.array([11, 8, 95, 51, 58, 21, 13]))
True
```


Tests for q3a (points: 0.5)
```
>>> bool(brand_cqrt_mean == 60.825)
True
```
```
>>> bool(brand_cqrt_mean - 60.825 < 0.0001)
True
```


Tests for q3b (points: 0.5)
```
>>> bool(brand_cqrt_median == 61)
True
```
```
>>> bool(brand_cqrt_median - 61 < 0.0001)
True
```


Tests for q3c (points: 0.5)
```
>>> bool(np.all(brand_total_minutes == np.array([3330, 1340, 256, 2352, 1320, 637, 348, 3729, 855, 2288, 1540, 4692, 1296, 2240, 795, 1212, 5232, 770, 552, 882, 888, 408, 3015, 540, 2507, 1025, 1800, 5175, 2880, 84, 2000, 361, 4485, 3300, 3763, 714, 2420, 2520, 2232, 2392, 252, 2820, 980, 1326, 120, 225, 4343, 3536, 522, 924, 1764, 2730, 2121, 3630, 1710, 850, 2562, 2200, 1550, 1856, 1207, 1295, 315, 1656, 648, 24, 1500, 1350, 3577, 585, 32, 559, 4876, 2272, 3036, 4264, 374, 252, 2944, 3440, 96, 3588, 2444, 2673, 1155, 2266, 4410, 2968, 1672, 980, 190, 3520, 4551, 2890, 144, 4738, 3103, 4200, 962, 1794, 4386, 1568, 1519, 504, 2940, 370, 1188, 330, 4800, 1462, 2867, 86, 1113, 2460, 525, 2376, 2220, 1368, 1302, 324])))
True
```
```
>>> len(brand_total_minutes) == 120
True
```
