```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment05_COMM187W26.ipynb")
```

# COMM 187 W26 Coding Assignment #5

<span style="color:red">**!!! PLEASE MAKE SURE TO RUN THE FIRST CELL ABOVE WHICH SAYS "# Initialize Otter" BEFORE RUNNING AUTO-GRADER !!!**</span>

## <center>Do Pulitzers Help Newspapers Keep Readers?</center>

![](./imgs/A5_introimg.png)

The dataset for this assignment is about the number of Pulitzer awards won by top 50 newspapers in the US (as of 2004), and their daily circulation. This dataset was collected and analyzed for the following [blog post by FiveThirtyEight](https://fivethirtyeight.com/features/do-pulitzers-help-newspapers-keep-readers/), and you can read more information about dataset [here](https://github.com/fivethirtyeight/data/tree/master/pulitzer).

The dataset is stored in the "data" subfolder and is named **`pulitzer-data.csv`**. It contains the following columns:


Header | Definition
--- | ---
`Newspaper` | Name of the newspaper
`Daily Circulation, 2004` | Daily circulation average as of 2004
`Daily Circulation, 2013` | Daily circulation average as of 2013
`Change in Daily Circulation, 2004-2013` | Percentage change in daily circulation average as of 2004 minus 2013
`Pulitzer Prize Winners and Finalists, 1990-2003` | Number of Pulitzer award winners and finalists between 1990 and 2003
`Pulitzer Prize Winners and Finalists, 2004-2014` | Number of Pulitzer award winners and finalists between 2004 and 2014
`Pulitzer Prize Winners and Finalists, 1990-2014` | Number of Pulitzer award winners and finalists between 1990 and 2014
---

```python
# RUN THIS CELL BEFORE ATTEMPTING THE ASSIGNMENT QUESTIONS
import numpy as np
import pandas as pd
from scipy import stats

df = pd.read_csv("./data/pulitzer-data.csv")
```

---

## Question 1

### A. Covariance between the daily circulations in 2004 and 2013 (0.5 point)

Calculate the covariance using the `.cov()` method in `pandas` between the daily circulations in 2004 and the daily circulations in 2013, and store it in the variable `q1_cov`.

Refer to the Week #4 Coding Lab for help on the `.cov()` method.

Note: Do not use the `np.cov()` function from `numpy`. 

```python
### Write your code below (in place of ...)
q1_cov = ...
```

```python
grader.check("q1a")
```

---

### B. Correlation between the daily circulations in 2004 and 2013 (0.5 point)

Calculate the correlation using the `.corr()` method in `pandas` between the daily circulations in 2004 and the daily circulations in 2013, and store it in the variable `q1_corr`.

Refer to the Week #4 Coding Lab for help on the `.corr()` method.

```python
### Write your code below (in place of ...)
q1_corr = ...
```

```python
grader.check("q1b")
```

---

## Question 2: Relation between Pulitzer Prize winners and change in circulation
### A: Correlation between number of Pulitzer award winners and finalists (1990 - 2003) and change in circulation (2004 - 2013) (0.5 point)

Does the number of Pulitzer prizes acquired by a newspaper during a decade correlate with the changes in circulation of the newspaper in the decade after?

Calculate the correlation using the `.corr()` method in `pandas` between the number of Pulitzer award winners and finalists (1990 - 2003) and the percentage change in circulation between 2004 and 2013, and store it in the variable `q2a`.

```python
### Write your code below (in place of ...)
q2a = ...
```

```python
grader.check("q2a")
```

---

### B: Correlation between number of Pulitzer award winners and finalists (2004 - 2014) and change in circulation (2004 - 2013) (0.5 point)

Does the number of Pulitzer prizes acquired by a newspaper during a decade correlate with the changes in circulation of the newspaper in the same decade?

Calculate the correlation using the `.corr()` method in `pandas` between the number of Pulitzer award winners and finalists (2004 - 2014) and the percentage change in circulation between 2004 and 2013, and store it in the variable `q2b`.

```python
### Write your code below (in place of ...)
q2b = ...
```

```python
grader.check("q2b")
```

---

### C: Correlation between number of Pulitzer award winners and finalists (1990 - 2014) and change in circulation (2004 - 2013) (0.5 point)

Do the changes in circulation of newspapers in a decade correlate with the number of Pulitzer prizes acquired by a newspaper during that decade as well as the decade before?

Calculate the correlation using the `.corr()` method in `pandas` between the number of Pulitzer award winners and finalists (1990 - 2014) and the percentage change in circulation between 2004 and 2013, and store it in the variable `q2c`.

```python
### Write your code below (in place of ...)
q2c = ...
```

```python
grader.check("q2c")
```

---

## Question 3: Analysis without high circulation newspapers

It is possible that correlations observed thus far are skewed, or inflated, because of the high circulation newspapers. High circulation newspapers are typically well resourced, and employ journalists who have high chances of winning the Pulitzer prize. 

Let's conduct an analysis without the high circulation newspapers, and notice if the correlation changes. 

**First step (A):** identify the newspapers for which the circulation is high, and make a new dataframe which excludes them.\
**Second step (B):** in this new dataset, calculate the correlation between the number of Pulitzer award winners and finalists (1990 - 2014) and change in circulation (2004 - 2013)

### A. Removing high circulation newspapers (1 point)

For this exercise, "high circulation newspapers" are defined as the ones for which the daily circulation in 2013 was higher than the mean + standard deviation of all circulations in 2013. (Can you think of a reason why we are choosing this definition of "high circulation newspapers"? Post your best guess in this week's Discussion Board.)

Identify these newspapers, and make a new dataframe `df_filtered` which does not include these high circulation newspapers.

**Steps:**
 - Calcualte the mean of "Daily Circulation, 2013" column using `.mean()` method and store it in a new variable `mean_circ_2013`
 - Calcualte the standard deviation of "Daily Circulation, 2013" column using `.std()` method and store it in a new variable `std_circ_2013`
 - Add the mean and standard deviation and store it in a new variable `mean_plus_std_2013`
 - Filter `df` to keep rows in which the value in "Daily Circulation, 2013" is less than or equal to `mean_plus_std_2013`, and store the result in a new variable `df_filtered` \
   Note: You have done a similar filtering in previous exercises, including Question 1 of Coding Assignment #4! You got this. Piece of cake.

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q3a")
```

---

### B. Correlation between the number of Pulitzer award winners and finalists (1990 - 2014) and change in circulation (2004 - 2013) in the newly filtered data (0.5 point)

Calculate the correlation using the `.corr()` method in `pandas` between the number of Pulitzer award winners and finalists (1990 - 2014) and the percentage change in circulation between 2004 and 2013 in `df_filtered`, and store it in the variable `q3b`.

Remember to use the newly created `df_filtered` and not `df` for this question!

```python
### Write your code below (in place of ...)
q3b = ...
```

```python
grader.check("q3b")
```

---

### C. Inference. (0 points)

This part is ungraded and just for your own practice. 

Compare the values of `q2c` and `q3b`. Both of these are correlations between the number of Pulitzer award winners and finalists (1990 - 2014) and change in circulation (2004 - 2013). `q2c` is the correlation when we **include** the high circulation newspapers, and `q3b` is the correlation when we **remove** the high circulation newspapers. 

Do you notice a difference? By just removing a few high circulation newspapers, what kind of difference are we able to observe in the correlations? What did you infer from this? **Post your inference on this week's Discussion Board.**

Tests for q1a (points: 0.5)
```
>>> bool(q1_cov - 161106010846.87592 < 0.0001)
True
```


Tests for q1b (points: 0.5)
```
>>> bool(q1_corr - 0.922249153863201 < 1e-05)
True
```


Tests for q2a (points: 0.5)
```
>>> bool(q2a - 0.2943850085162899 < 1e-05)
True
```


Tests for q2b (points: 0.5)
```
>>> bool(q2b - 0.4025279361189322 < 1e-05)
True
```


Tests for q2c (points: 0.5)
```
>>> bool(q2c - 0.3523112737213655 < 1e-05)
True
```


Tests for q3a (points: 1)
```
>>> df_filtered.shape == (47, 7)
True
```
```
>>> not any(df_filtered['Newspaper'].isin(['USA Today', 'Wall Street Journal', 'New York Times']))
True
```
```
>>> bool(df_filtered['Daily Circulation, 2013'].max() == 653868)
True
```
```
>>> bool(df_filtered['Daily Circulation, 2004'].max() == 983727)
True
```


Tests for q3b (points: 0.5)
```
>>> bool(q3b - -0.0037546253012713765 < 0.0001)
True
```
