```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment02_COMM187W26.ipynb")
```

# COMM 187 W26 Coding Assignment #2

<span style="color:red">**!!! PLEASE MAKE SURE TO RUN THE FIRST CELL ABOVE WHICH SAYS "# Initialize Otter" BEFORE RUNNING AUTO-GRADER !!!**</span>

### Question 1

Consider the following Python dictionary `productdata`, with product names as its keys, and its values are dictionaries that contain more information about the product.

#### A. (0.5 point)

Add a new product "T-shirt" with the following information to the dictionary `productdata`:
 - brand: Goodfellows
 - product_type: Clothing
 - price: 22.99

To do so, add a new key:value pair to `productdata`, where key is "T-shirt" and value is a dictionary with the product information.

```python
productdata = {"Wireless Headphones": {"brand": "SoundCore", "product_type": "Electronics", "price": 149.99},
               "Coffee Maker": {"brand": "BrewMaster", "product_type": "Kitchen Appliance", "price": 89.95},
               "Running Shoes": {"brand": "Velocity", "product_type": "Footwear", "price": 129.50},
               "Organic Shampoo": {"brand": "NaturalBliss", "product_type": "Personal Care", "price": 18.75},
               "Desk Lamp": {"brand": "LuminaHome", "product_type": "Home Décor", "price": 45.25},
               "Protein Powder": {"brand": "FitFuel", "product_type": "Nutrition", "price": 32.99}}

### Write your code below (in place of ...)
...
```

```python
grader.check("q1a")
```

#### B. (0.25 point)

Convert `productdata` to a pandas DataFrame `productdf`. It should look like this:

![](./imgs/A2_Q1B_img.png)

**Tip:** Simply use the DataFrame function in pandas to convert the Python dictionary to a pandas DataFrame. 

```python
import pandas as pd
### Write your code below (in place of ...)
productdf = ...
```

```python
grader.check("q1b")
```

### Question 2

#### A. (0.5 point)

Load the CSV file named `recent-grads-mod.csv` from the `data` folder and store it in a variable called `grads_df`. \
Use the function `pd.read_csv` as described in this week's coding lab and Datacamp module.

Context: This dataset is a modified version of the data behind [this article](https://fivethirtyeight.com/features/the-economic-guide-to-picking-a-college-major/). This data shows the earnings of Americans with different college majors from a survey conducted in 2014.

```python
### Write your code below (in place of ...)
grads_df = ...
```

```python
grader.check("q2a")
```

#### B. (0.5 point)

Use `.iloc()` to select the rows corresponding the majors "COMMUNICATION" (at index 92) and "MASS MEDIA" (at index 124) and store it in a new pandas DataFrame `comm_media`.

```python
### Write your code below (in place of ...)
comm_media = ...
```

```python
grader.check("q2b")
```

#### C. (0.25 point)

Extract the `Total` column as a pandas Series and store it in a variable called `total`.

```python
### Write your code below (in place of ...)
total = ...
```

```python
grader.check("q2c")
```

#### D. (0.5 point)

Create a boolean pandas Series called `total_greater_than_5k` that contains `True` for each value in the pandas Series `total` that is greater than five thousand, and `False` for all other values.

```python
### Write your code below (in place of ...)
total_greater_than_5k = ...
```

```python
grader.check("q2d")
```

#### E. (0.5 point)

Make a new pandas DataFrame called `grads_df_large` with the data for only those majors for which the total number of majors were greater than five thousand. 

Steps:
 - Use the boolean pandas Series `total_greater_than_5k` to subset the `grads_df` DataFrame.
 - Store the resulting selection in `grads_df_large`.

```python
### Write your code below (in place of ...)
grades_df_large = ...
```

```python
grader.check("q2e")
```

#### F. (1 point)

Make a new pandas DataFrame called `grads_df_parttime_20k` with the data for only those majors for which the number of partime workers, which is stored in the column `Part_time`, is greater than or equal to 20 thousand. 

Follow the steps similar to the ones used in the Question 2 C through E:
 - Extract the `Part_time` column as a pandas Series and store it in a variable called `parttime`.
 - Create a boolean pandas Series called `parttime_greater_than_20k` that contains `True` for each value in the pandas Series `parttime` that is greater than or equal to 20,000, and `False` for all other values.
 - Use the boolean pandas Series `parttime_greater_than_20k` to subset the `grads_df` DataFrame.
 - Store the resulting selection in `grads_df_parttime_20k`. This is the only variable that will be checked for accuracy by the autograder.

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q2f")
```

Tests for q1a (points: 0.5)
```
>>> 'T-shirt' in productdata.keys()
True
```
```
>>> productdata['T-shirt']['brand'] == 'Goodfellows'
True
```
```
>>> productdata['T-shirt']['product_type'] == 'Clothing'
True
```
```
>>> productdata['T-shirt']['price'] == 22.99
True
```


Tests for q1b (points: 0.25)
```
>>> all(productdf == pd.DataFrame(productdata))
True
```


Tests for q2a (points: 0.5)
```
>>> all(grads_df == pd.read_csv('./data/recent-grads-mod.csv'))
True
```


Tests for q2b (points: 0.5)
```
>>> type(comm_media) == pd.core.frame.DataFrame
True
```
```
>>> all(comm_media == grads_df.iloc[[92, 124]])
True
```


Tests for q2c (points: 0.25)
```
>>> type(total) == pd.core.series.Series
True
```
```
>>> all(total == grads_df['Total'])
True
```


Tests for q2d (points: 0.5)
```
>>> type(total_greater_than_5k) == pd.core.series.Series
True
```
```
>>> sum(total_greater_than_5k) == 125
True
```
```
>>> all(total_greater_than_5k == (total > 5000))
True
```


Tests for q2e (points: 0.5)
```
>>> grades_df_large.shape == (125, 21)
True
```
```
>>> all(grades_df_large == grads_df[total_greater_than_5k])
True
```


Tests for q2f (points: 1)
```
>>> grads_df_parttime_20k.shape == (21, 21)
True
```
```
>>> all(grads_df_parttime_20k == grads_df[grads_df['Part_time'] >= 20000])
True
```
