```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment04_COMM187W26.ipynb")
```

# COMM 187 W26 Coding Assignment #4

<span style="color:red">**!!! PLEASE MAKE SURE TO RUN THE FIRST CELL ABOVE WHICH SAYS "# Initialize Otter" BEFORE RUNNING AUTO-GRADER !!!**</span>

## <center>TWEETS BY US SENATORS (2008-17)</center>

![](./imgs/A3_introimg.png)

The dataset for this assignment is about the favorites (likes), replies, and retweets on tweets made by US Senators between 2008 and 2017, as recorded on Oct. 19 and 20, 2017. This dataset was collected and analyzed for the following [blog post by FiveThirtyEight](https://fivethirtyeight.com/features/the-worst-tweeter-in-politics-isnt-trump/), and you can read more information about dataset [here](https://github.com/fivethirtyeight/data/tree/master/twitter-ratio/).

The original dataset for this article contains over 288 thousand rows. Unlike last week, this week you have beeen provided with a reduced dataset of randomly selected **forty thousand rows**, which is stored in the "data" subfolder and is named **`senators_40k.csv`**. It contains the following columns:


Header | Definition
--- | ---
`created_at` | Date and time of the creation of the tweet in the format "mm/dd/yy hh:mm"
`text` | The text content of the tweet
`url` | URL of the tweet
`replies` | Number of replies for the tweet when data was collected
`retweets` | Number of retweets for the tweet when data was collected
`favorites` | Number of favorites (or likes) for the tweet when data was collected
`user` | Twitter username of the US Senator
`bioguide_id` | Unique ID for each US Senator
`party` | Party identification for each US Senator: 'D' for Democrat, 'R' for Republican, and 'I' for Independent
`state` | The state being represented by the US Senator

---

```python
# RUN THIS CELL BEFORE ATTEMPTING THE ASSIGNMENT QUESTIONS
import numpy as np
import pandas as pd
from scipy import stats

df = pd.read_csv("./data/senators_40k.csv")
```

---

## Question 1: Remove tweets with 0 favorites (1 point)

Remove rows from `df` with 0 favorites and store the resulting dataframe in `df_filtered`.

Steps:
 - Identify row indices for which the `favorites` column does not have value 0. You can do this by using the `!=` operator, which means "not equal to". 
 - Use these indices to extract rows with `favorites` not equal to 0 and store it in a new variable `df_filtered`.
 - You have done similar filtering in previous exercises, including the hackathon! You got this.

```python
### Write your code below (in place of ...)
df_filtered = ...
```

```python
grader.check("q1")
```

---

## Question 2: Create a new column `ratio` in `df_filtered` (1 point)

As you might be aware, **"ratio"** on social media refers to the ratio of *replies or comments* to *likes* (as defined [here](https://www.howtogeek.com/721651/what-does-ratio-mean-on-social-media/)). On Twitter, a ratio could mean the ratio of either repliess or retweets to favorites. For this question, we will define "ratio" as the ratio of replies to favorites.

It has been observed that high ratio of replies to likes suggests that a post is unpopular or controversial. The reasoning seems to be that if more people are choosing to reply or comment instead of liking the post, they did so to express their disapproval.

Make a new column in `df_filtered` called `ratio` which is calculated by dividing the value in `replies` by the value in `favorites` for each tweet.

**Steps:**
- Make a new variable `ratio` which is the division of `replies` column with `favorites` column.
- Make a new column called `ratio` in `df_filtered`, and assign it the value of the variable `ratio`. You might get a warning for this, ignore it.
- Notice that the name of the new column as well as the new variable are the same -- `ratio`. Notice the difference, and avoid any confusion!

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q2")
```

---

## Question 3: Identify the most ratio'ed tweet in `df_filtered` (1 point)

Identify the most ratio'ed tweet in `df_filtered` and store it in `most_ratioed`.

**Steps:**
- Use the `.idxmax()` method to identify the index of the maximum value in the `ratio` column and store it in a new variable `max_ratio_index`. This is a new method which we have not covered in class. It is very simple to use, and you are fully capable to use it based on [this documentation](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.idxmax.html). Try it, and if you need any help, reach out to your group or post on the Discussion Board.
- Extract the tweet text for this index using `.loc[max_ratio_index, 'text']` and store it in `most_ratioed`.
- Follow these steps and don't simply copy the text of the most ratioed tweet here. If you do, it will be identified later and points will be deducted.

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q3")
```

---

## Question 4: Comparing Ratios on Tweets Across Parties in `df_filtered` (1 point)

Compare the ratios across the three parties (stored in the column `party` in `df_filtered` with values **D**, **R**, and **I**) using ANOVA. \
Store the F statistic in `q4_f` and the p-value in `q4_p`.

Follow similar steps as in Question 3A from the previous assignment. You are not being provided with a list of explicit steps for this question, but you have had enough practice to make this work. You got this :)

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q4")
```

---

Tests for q1 (points: 1)
```
>>> df_filtered.shape == (36498, 10)
True
```
```
>>> all(df_filtered == df[df['favorites'] != 0])
True
```


Tests for q2 (points: 1)
```
>>> all(df_filtered['ratio'] == df_filtered['replies'] / df_filtered['favorites'])
True
```


Tests for q3 (points: 1)
```
>>> most_ratioed == "MT @ncmbc Don't forget to register: NC Aerospace Suppliers Conference, August 24. @SenatorBurr will kick off this event http://t.co/9Yc4Yne"
True
```


Tests for q4 (points: 1)
```
>>> bool(abs(q4_f - 937.5702574223225 < 0.001))
True
```
```
>>> bool(abs(q4_p - 0 < 0.001))
True
```
