```python
# Initialize Otter
import otter
grader = otter.Notebook("DemoAssignCOMM187W26.ipynb")
```

# COMM187: Data Science in Communication Research, Winter 2026 -- Demo-Assignment

**Question 1:** Write a statement which assigns `2*3` to a variable named `var`.

```python
### Write your code below (in place of ...)
var = ...
```

```python
grader.check("q1")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

This is a DEMO assignment for COMM187: Data Science in Communication Research. Please answer the coding question below, as instructed. Then, download this .ipynb file, and submit it on Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 1)
```
>>> var == 6
True
```
